# Linear Regression From Scratch
In this notebook you will implement linear regression using the least squares method.

The notebook will guide you in the general steps. You may use only numpy, pandas, and matplotlib for this exercise.

#### Math Background
The task is to solve linear regression for the data set of ```x``` and ```y```. That is, fit a line y to the data such that ```y = a + bx + e```. Where a and b are coefficents and e is an error term.

We know that ```b = SUM ( xi - x_mean) * (yi - y_mean) / SUM (xi - x_mean)^2``` where ```xi```, and ```yi``` are the indivdual observations and ```x_mean```, ```y_mean``` are means of all ```xi``` and ```yi```.

So a reasonable solution then for a fit is ```a = y_mean - b * x_mean```.


#### Implementation Steps
1. Load the file ```reg_data.csv```
2. Display the data with a scatter plot. Write a markdown cell and comment on the structure of the data and what kind of line you think will fit the data. 
3. Implement a function of the least squares method.
4. Plot the predictions from your function.
5. Comment on the output of your function. Does it make sense? Can it be made better? If yes, how?

```python
#import the minimum packages
import numpy as np
import pandas as pd
import matplotlib.pyplot as plt
from sklearn.model_selection import train_test_split
import os
```

<!-- grade_id: cell-4104816f50e4844b, points: 0, grade: false, solution: true -->
```python
#load and inspect the data

# YOUR CODE HERE
for file in os.listdir():
    if file.endswith(".csv"):
        df = pd.read_csv(file)

df.head()

```

<!-- grade_id: cell-3676b87d8fc79bbd, points: 1, grade: true, solution: false -->
```python
assert df.shape == (100,2)
```

<!-- grade_id: cell-7f576e50392bc5e9, points: 0, grade: false, solution: true -->
```python
#plot the data (scatter)
# YOUR CODE HERE
X = df['X']
Y = df['Y']
plt.figure(figsize=(10,6))
plt.scatter(X, Y, label = "Data" )
plt.xlabel('X')
plt.ylabel('Y')
```

<!-- grade_id: cell-54891567ed3179cb, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
x_train, x_test, y_train, y_test = train_test_split(X, Y, test_size=0.2, random_state = 0)
```

<!-- grade_id: cell-161d857ab77810d3, points: 1, grade: true, solution: false -->
```python
assert x_train.shape == (80,)
assert x_test.shape == (20,)
assert y_train.shape == (80,)
assert y_test.shape == (20,)
```

<!-- grade_id: cell-d5e750d407949b20, points: 0, grade: false, solution: true -->
```python
#implement a least squares function to find a, b
# mean(x)*mean(y) - mean(x*y)/((mean(x)^2) - mean(x^2))

x_mean = 0
y_mean = 0
b_upper = 0
b_lower = 0
b = 0
a = 0

# YOUR CODE HERE
def least_squares(x,y):
    x_mean = np.mean(x)
    y_mean = np.mean(y)
    xy_mean = np.mean(np.multiply(x, y))
    x2_mean = np.mean(np.multiply(x,x))
    b = ((x_mean*y_mean) - xy_mean)/((x_mean**2) - x2_mean)
    # a + b*x_mean = y_mean
    # a = y_mean - b*x_mean 
    a = y_mean - (b*x_mean)
    return a, b

a, b = least_squares(x_train, y_train)
print(a,b)
```

<!-- grade_id: cell-c7fea6e45dc92701, points: 5, grade: true, solution: false -->
```python
assert abs(a - 6.68) <=0.03
assert abs(b - 1.34) <=0.03
```

<!-- grade_id: cell-721384e1a2152077, points: 0, grade: false, solution: true -->
```python
line = 0
x = np.array([])
y = np.array([])
# YOUR CODE HERE
x = np.append(x, X)
y = np.append(y, (a + b*x))
plt.plot(x, y, 'r', label = 'regression line')
plt.scatter(X, Y, label = 'data')
plt.xlabel('X')
plt.ylabel('Y')
plt.legend(loc="upper left")
```

<!-- grade_id: cell-71f1201b79ff7e44, points: 5, grade: true, solution: false -->
```python
assert abs(a +3*b - 10.72) <= 0.05
assert abs(a +100*b - 140) <= 2
```

<!-- grade_id: cell-8b80d3f35ea3b407, points: 0, grade: false, solution: true -->
```python
#Separate your test data in to classes
#if the Y value for a certain X is lower than the line then the class is 0
class_0 = []
class_1 = []

# YOUR CODE HERE
for x,y in zip(x_test,y_test):
    if y < (x*b + a):
        class_0.append((x,y))
    else:
        class_1.append((x,y))

class_0 = np.array(class_0)
class_1 = np.array(class_1)
print(class_0.shape)
print(class_1.shape)
```

<!-- grade_id: cell-73e31e3494663d82, points: 5, grade: true, solution: false -->
```python
assert 9 < class_0.shape[0] < 13
assert 7 < class_1.shape[0] < 11
```

<!-- grade_id: cell-db81fd9fa3a21bdf, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
x = np.linspace(30, 70, 100)
plt.plot(x, x*b + a, linestyle = '-', label = 'regression line')
plt.scatter(class_0[:,0], class_0[:,1], label = 'Class 0', c = 'red')
plt.scatter(class_1[:,0], class_1[:,1], label = 'Class 1', c = 'green')
plt.xlabel('x')
plt.ylabel('y')
plt.legend()
```

<!-- grade_id: cell-ee8df562ed62329a, points: 0, grade: false, solution: true -->
```python
# get the total error for the classes

err = 0
# YOUR CODE HERE
classes = np.concatenate((class_0,class_1), axis = 0)
for x,y in classes:
    err += abs(y - (x*b + a))
```

<!-- grade_id: cell-bbe692ec61946f5c, points: 5, grade: true, solution: false -->
```python
assert err - 201 < 20
```

```python

```